# Lab 2: Numpy

*Since this is Numpy exercise, you are not allowed to use loop and I will tell you which part you can*

(Last update: 07/11/2022)



Name: Lê Nguyễn Bình Nam

Student code: 20127567

---

## Instructions for doing and submitting assignment

&#9889; Note that I will use a software for grading your assignment so you have to follow exactly what I said. If you are not sure about something, please raise your question on Zalo group.

**How to do your assignment**

You will do your assignment directly on this notebook file. First, you fill your name and student code at the beginning of the file. In this file, you will write your code when you see the following lines of code:
```python
# YOUR CODE HERE
raise NotImplementedError()
```

For optional coding parts, there will be:
```python
# YOUR CODE HERE (OPTION)
```

For markdown cell, there will be:
```markdown
YOUR ANSWER HERE
```

Of course, you have to remove the `raise NotImplementedError()` statement when you finish.

For coding parts, there are often cells below to help you check your answers. You will pass the test if there are no errors when you run the test cells. In some cases, the tests are insufficient. That means that if you do not pass the test, your answer is definitely wrong somewhere, but if you pass the test, your answer may still be incorrect.
 
While doing the assignment, you should print out the output and create more cells for testing. But you have to remove all of them (comment your print-out codes, delete the cell created by you) when you submit your code. <font color=red>Do not remove or edit my cells</font> (except for the aforementioned cells).

Keep your code clean and clear by using meaningful variable names and comments, not write too-long coding lines.
Press `Ctrl + S` right after editing.

Keep it real: The reason why you are here is to <font color=green>study, really study</font>. I highly recommend that you discuss your idea with your friends and <font color=green>write your own code based on your own knowledge</font>. <font color=red>Copy means zero.</font>

**How to submit your assignment**

When grading your assignment, first, I will choose `Kernel` - `Restart & Run All` in order to restart the kernel and run all cells in your notebook. Therefore, you should do that before submitting to ensure that the outputs are all as expected.

After that, you make a submited direction as follow:
- Folder `Student code` (for example, if your student code is 1234567, then you name your folder `1234567`)
    - File `Lab2.ipynb` only

Finally, you compress your folder (`Student code`) and submit on Moodle. <font color=red>Please strictly follow the submission rules.</font>

---

## Programming environment

- You will re-use the Linux environment setup in Lab 0 - WarmUp. Don't forget to start your coding environment (`conda activate min_ds-env`) before doing your assignment.
- Use Jupyter notebook or Jupyter lab, <font color=red>not Google Colab</font> (I can not grade you well on Google Colab) to edit your `*.ipynb` file.

```python
import sys
sys.executable
```

Output:
```
'/root/miniconda3/envs/min_ds-env/bin/python'
```

Your file now should run on `min_ds-env`.

You will not use `pandas` here.

---

## Import necessary libraries

```python
import numpy as np
# YOUR CODE HERE (OPTION)
```

---

## Data collection

In this assignment, we simply use a collected dataset. You can download it [here](https://files.grouplens.org/datasets/movielens/ml-100k/u.data) and save the file as "u.data" (not "u.txt" or anything else). The dataset has more than one file but you just have to use a single file named "u.data".

According to the [data description file](https://files.grouplens.org/datasets/movielens/ml-100k-README.txt):
> MovieLens data sets were collected by the GroupLens Research Project
at the University of Minnesota.
> 
> This data set consists of:
> * 100,000 ratings (1-5) from 943 users on 1682 movies. 
> * Each user has rated at least 20 movies.  
> * Simple demographic info for the users (age, gender, occupation, zip)
>
> The data was collected through the MovieLens web site
(movielens.umn.edu) during the seven-month period from September 19th, 
1997 through April 22nd, 1998. This data has been cleaned up - users
who had less than 20 ratings or did not have complete demographic
information were removed from this data set. Detailed descriptions of
the data file can be found at the end of this file.

The above is the general description of all file in dataset. The description of "u.data" is in the following:
> The full u data set, 100000 ratings by 943 users on 1682 items. Each user has rated at least 20 movies. Users and items are numbered consecutively from 1. The data is randomly ordered. This is a tab separated list of user id | item id | rating | timestamp. The time stamps are unix seconds since 1/1/1970 UTC 

---

## Data exploration

<!-- ### Dữ liệu có bao nhiêu dòng và bao nhiêu cột? (0.5đ) -->
### How many rows and columns in dataset? (0.5p)

First of all, you have to read the file "u.data" into a numpy array named `raw_ratings` (use function `np.genfromtxt`). Let `dtype` be `np.uint64`. You also have to put "u.data" in the same directory with this notebook file so you can just pass "u.data" to the function.

Next, you calculate the number of rows and columns and save them into `n_rows` and `n_cols`.

<!-- grade_id: cell-1bf7d4be2f1eb278, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
dataset_path = './u.data'
raw_ratings = np.genfromtxt(dataset_path,
                        delimiter="\t",
                        dtype=np.uint64)

n_rows = raw_ratings.shape[0]
n_cols = raw_ratings.shape[1]
```

<!-- grade_id: cell-fd69b593a2722fbc, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert raw_ratings.dtype == np.uint64
assert raw_ratings.ndim == 2
assert n_rows == 100000
assert n_cols == 4
raw_ratings[:5] # Xem 5 dòng đầu tiên
```

Output:
```
array([[      196,       242,         3, 881250949],
       [      186,       302,         3, 891717742],
       [       22,       377,         1, 878887116],
       [      244,        51,         2, 880606923],
       [      166,       346,         1, 886397596]], dtype=uint64)
```

### What is the meaning of each row? Are there any 2 rows with different meaning?

According to data description file and by taking a quick look at the data, each line provides information about the score that a user rates a movie. It seems that there is no outlier line.

<!-- ### Dữ liệu có các dòng bị lặp không? (0.5đ) -->
### Is the data duplicated? (0.5p)

<!-- Bạn sẽ kiểm tra vụ này và lưu kết quả vào biến `have_duplicated_rows`. Biến này sẽ có giá trị True nếu dữ liệu có các dòng bị lặp và có giá trị False nếu ngược. -->
You are going to check that and save the result into `have_duplicated_rows`. The value of this variable is `True` if the data is not duplicated and vice versa. 

<!-- grade_id: cell-c184a3a04425a1c1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
have_duplicated_rows = len(raw_ratings) != len(np.unique(raw_ratings, axis=0))
```

<!-- grade_id: cell-7f2d2f011cc33f7d, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert have_duplicated_rows == False
```

<!-- ### Mỗi cột có ý nghĩa gì? -->
### What is the meaning of each row?

According to the data description file:
- Col #1: The ID of the user who participated in the rating
- Col #2: The ID of the movie
- Col #3: The score that user rated
- Col #4: The time when user did the rating (in second, count from 1/1/1970 UCT)

<!-- ### Mỗi cột hiện đang có kiểu dữ liệu gì? Có cột nào có kiểu dữ liệu chưa phù hợp để có thể xử lý tiếp không? -->

### What is the datatype of each column? Are there any columns with datatype that is not suitable for further processing?

<!-- Hiện tại thì tất cả các cột đều có kiểu dữ liệu là số: -->
The datatypes of all columns are all numeric:

```python
raw_ratings.dtype
```

Output:
```
dtype('uint64')
```

All the columns have the suitable datatype.

You might think that the the datatype of the final column is not suitable? In fact, this column represents dates in the form of number. We still can find min/max values and can convert them to datetime format when needed. The point is that you can not store 2 different datatypes in the same numpy array. Therefore, if you change the datatype of that column, you have to prepare another numpy array to store these datetime values.

<!-- ### Với mỗi cột có kiểu dữ liệu dạng số (numeric) và dạng datetime, các giá trị được phân bố như thế nào? (1đ) -->
### How does the numeric/datetime data distribute? (1p)

It looks like that the datatype of columns are numeric. But you might notice that the datatype of "user id" column and "movie id" column should be catergorical. According to this approach, there will be 2 numeric columns ("rating" and "timestamp").

How many missing values do these two columns have?

```python
np.sum(np.isnan(raw_ratings[:, 2:]), axis=0)
```

Output:
```
array([0, 0])
```

What you are going to do next is to find the min, median, max values of these 2 numeric columns (use function `np.percentile`, let the `interpolation` option be `linear`). Then, save the result into 2 variables named `rate_col_profile` and `rate_date_col_profile`. These are 2 1d numpy arrays. The datatype of `rate_col_profile` is real number and the one of `rate_date_col_profile` is `datetime64[s]`

Hint: When converting an array, you can pass the datatype in the form of string.

<!-- grade_id: cell-15364ea0e0e9c792, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
numeric_ratings = raw_ratings[:, 2]
datetime_ratings = raw_ratings[:, -1].astype('datetime64[s]')

rate_col_profile = []
rate_col_profile.append(float(np.min(numeric_ratings)))
rate_col_profile.append(np.percentile(numeric_ratings, 50))
rate_col_profile.append(float(np.max(numeric_ratings)))

rate_date_col_profile = []
rate_date_col_profile.append(np.min(datetime_ratings))
rate_date_col_profile.append(np.percentile(datetime_ratings, 50))
rate_date_col_profile.append(np.max(datetime_ratings))
```

<!-- grade_id: cell-3f3b0054dd6eadb9, points: 1, grade: true, solution: false -->
```python
# TEST
assert np.array_equal(rate_col_profile, np.array([1., 4., 5.]))
assert np.array_equal(rate_date_col_profile, 
                      np.array(['1997-09-20T03:05:10', 
                                '1997-12-22T21:42:24',
                                '1998-04-22T23:10:38'], 
                               dtype='datetime64[s]'))
```

Min value and max value of these columns are normal and match the information in data description file.

<!-- ### Với mỗi cột có kiểu dữ liệu dạng phân loại (categorical), các giá trị được phân bố như thế nào? (1đ) -->
### How does the catergorical data distribute? (1p)

As aforementioned, the 2 first columns contain categorical data.

How many missing values do these two columns have?

```python
np.sum(np.isnan(raw_ratings[:, :2]), axis=0)
```

Output:
```
array([0, 0])
```

Now, for each column, you have to find out 5 values: number of different values; least occurrence value and number of occurrences (2 values); most occurrence value and number of occurrences (2 values); Save these list into `user_col_profile` and `movie_col_profile`.

Note: If there are more than one value with the least/most occurrence, just take the smallest/biggest value (hint: if you let `kind=stable` in sort function, Numpy will ensure that the order of the same values in the original array does not change).

<!-- grade_id: cell-09819763bd956c3a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
user = raw_ratings[:, 0].astype(np.int64)
num_users = len(np.unique(user))

counts = np.bincount(user)
min_value = min(counts[1:])
id_min = np.where(counts == min_value)[0]
id_min = id_min[0]
max_value = max(counts[1:])
id_max = np.where(counts == max_value)[0]
id_max = id_max[0]

user_col_profile = [num_users, id_min, min_value, id_max, max_value]

movie = raw_ratings[:, 1].astype(np.int64)
num_movies = len(np.unique(movie))

counts = np.bincount(movie)

min_value = min(counts[1:])
id_min = np.where(counts == min_value)[0]
id_min = id_min[0]

max_value = max(counts[1:])
id_max = np.where(counts == max_value)[0]
id_max = id_max[0]
movie_col_profile = [num_movies, id_min, min_value, id_max, max_value]
```

<!-- grade_id: cell-9f4bbdb1bcb99063, points: 1, grade: true, solution: false -->
```python
assert user_col_profile==[943, # Có chừng này user
                          19,  # Đây là user rate ít movie nhất
                          20,  # và đó là chừng này movie
                          405, # Đây là user rate nhiều movie nhất
                          737] # và đó là chừng này movie
assert movie_col_profile==[1682,#Có chừng này movie
                           599, #Đây là movie được ít user rate nhất
                           1,   #và đó là chừng này user
                           50,  #Đây là movie được nhiều user rate nhất
                           583] #và đó là chừng này user
```

These informations match the information in the data description file:
> The full u data set, 100000 ratings by 943 users on 1682 items. Each user has rated at least 20 movies. Users and items are numbered consecutively from 1.

The state "Users and items are numbered consecutively from 1" is not varified. You can check that by running the following code cell:

```python
print('User id  - min & max:', 
      raw_ratings[:, 0].min(), '&', raw_ratings[:, 0].max()) 
print('Movie id - min & max:', 
      raw_ratings[:, 1].min(), '&', raw_ratings[:, 1].max()) 
```

Output:
```
User id  - min & max: 1 & 943
Movie id - min & max: 1 & 1682
```

---

## Propose meaningful questions & Answer them

After exploration process, you already have a deeper understanding about the data. For now, you will propose meaningful questions that can answer by using the data. 

They could be: \
*Which movies should we recommend for each user?*

<!-- Việc tìm ra câu trả lời cho câu hỏi này sẽ giúp ích cho cả người dùng và nhà cung cấp dịch vụ xem phim: -->
All users and film providers will benefit a lot after answering this question:

<!-- - Người dùng: người dùng có thể muốn xem phim, nhưng có rất nhiều phim và người dùng không biết nên xem phim nào; sẽ rất tiện cho người dùng nếu hệ thống có thể đề xuất một danh sách phim mà nhiều khả năng người dùng sẽ thích -->
- Users: There are a lot of movies and users might not know which movies to choose. It is so great if your system can provide a list of movies that users may love them.
<!-- - Nhà cung cấp dịch vụ xem phim: nếu hệ thống đề xuất tốt thì nhiều khả năng người dùng sẽ xem phim và sẽ thích, nghĩa là người dùng sẽ tiếp tục trả phí để sử dụng dịch vụ ;-) -->
- Film provider: If your system work well, users will keep using your service and pay more.

---

## Pre-processing (1p)

<!-- Việc xây dựng hệ thống đề xuất sản phẩm mình đã có nói sơ bộ ở buổi học 09. Đầu tiên, từ dữ liệu ban đầu là array `raw_ratings`, bạn sẽ cần tạo array `ratings` gồm số-lượng-user dòng và số-lượng-movie cột, `ratings[i, j]` cho biết user có id i đã cho movie có id j bao nhiêu điểm, nếu chưa cho thì giá trị sẽ là `nan`. Vì Python đánh chỉ số từ 0 nên bạn cũng cần chỉnh lại để id của user và movie bắt đầu từ 0 (thay vì 1). -->

First, you have to create array `ratings` from `raw_ratings`. `ratings` is a 2d numpy array. `rattings[i,j]` represents the score that user i graded movie j. If user i does not grade movie j, then `ratings[i,j]=nan`.

Note: The index should start from 0. That means user 1 in `ratings` will be represented by 0.

<!-- grade_id: cell-11ff284c3fe7fb7e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
user = raw_ratings[:,0] - 1
movie = raw_ratings[:,1] - 1

row = np.unique(user).size
col = np.unique(movie).size

ratings = np.zeros((row, col))

ratings[user, movie] = raw_ratings[:, 2]
ratings[ratings == 0] = np.nan
```

<!-- grade_id: cell-fb18066caaa407ed, points: 1, grade: true, solution: false -->
```python
assert ratings.shape == (943, 1682)
missing_ratios = np.mean(np.isnan(ratings))
assert missing_ratios.round(3) == 0.937 # Có 93.7% giá trị nan!
```

---

## Analyze data to answer the question (6p)

<!-- Để có thể đề xuất các movie cho mỗi user thì ta sẽ phải làm 2 công đoạn: -->
In order for your system to recommend movies, you have to do 2 steps:

- Step #1: Predict the score that a user grade an unseen movie. In the other words, you have to fill in the nan values in array `ratings`.

- Step #2: Recommend unseen movies that have high score for user.

In this assignment, we only forcus on step #1.

I have uploaded some materials that can help you learn how to use NumPy better as well as how to predict movie scores for a user. In this assignment, you will predict the scores for all users. Note that you are not allowed to use loops (for, while) since this is the lab about Numpy. Therefore, instead of doing that, you will predict the scores for all users at the same time. In order for you to do that, you will have to predict the score that a user will give not just a movie but all movies (even movies that he or  she has already watched). Calculating scores for movies that he/she watched is not necessary, but it will help your calculation process speed up (because we calculate for all users at the same time).

In fact, you can't process all users at the same time because of a lack of memory (the temporary arrays generated in the computing process will have the size of `#users * #users * #movies`). In order to solve that problem, you can divide your data into batches (a group of users) so the temporary arrays will only have the size of `batch_size * #users * #movies`. When you successfully calculate the scores for a batch of users, you can use loops (for, while) to calculate the scores for all users by traversing all batches.

Let `batch_size=32`, you will have to fill the array `filled_ratings` with the scores of all movies. When doing your homework, **you should not change the `ratings` array**.

```python
batch_size = 32
filled_ratings = np.empty_like(ratings)
```

I do not mean to scare you, but the implementation is quite complex and mainly based on **broadcasting mechanism**. Therefore, first, you should try predicting the score on a batch. I will help you on that by dividing it into steps and providing test suits for each of them. The important point is that somehow, you change from predicting scores for a user to predicting scores for a batch of users: there will be a dimension at the beginning of the array corresponding to the size of `batch_size`.

First, you will predict the scores for a batch of users indexing from `start` to `end`.

```python
start = 0
end = batch_size
```

**Step #1**: Calculate `similarities` array that shows the similarity between each client in the batch under consideration and all uses. The size of this array will be `batch_size * #users`. `similarities[i,j]` represents the similarity level of user i (in the batch) and user j (in all users). If there are 2 users that do not overlap any movies, there will be a warning that says "RuntimeWarning: Mean of empty slice." In this case, just let the similarity be 0.

<!-- grade_id: cell-1d34779ee4335381, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
first_batch = ratings[start:end]
similarities = np.nanmean(np.abs(np.expand_dims(ratings, 0) - np.expand_dims(first_batch, 1)), axis=2)
similarities = 1 / (similarities + 0.001)
similarities[np.isnan(similarities)] = 0.0
```

Output:
```
/tmp/ipykernel_859/147193662.py:3: RuntimeWarning: Mean of empty slice
  similarities = np.nanmean(np.abs(np.expand_dims(ratings, 0) - np.expand_dims(first_batch, 1)), axis=2)
```

<!-- grade_id: cell-3c3f855070577541, points: 2, grade: true, solution: false -->
```python
# TEST
assert similarities.shape == (32, 943)
assert np.array_equal(similarities[:3, :3].round(1), 
                      np.array([[1.0e+03, 1.1e+00, 7.0e-01],
                                [1.1e+00, 1.0e+03, 7.0e-01],
                                [7.0e-01, 7.0e-01, 1.0e+03]]))
```

**Step #2**: Compute the array "weights" as a normalized version of array "similarity". For each user in the batch, the similarity with all other users will be normalized for each movie: normalize for users who have given movie scores by dividing the similarity by the total similarity that only considers users who have given movie scores, with the rest of the users normalizing by giving similarity zero. The array "weights" will have dimensions of `batch_size * #users * n_movies`. When running, you will see a warning: `RuntimeWarning: invalid value encountered in true_divide.` This is because the users that rate the movie have a similarity of 0 for all users, resulting in a normalization of 0/0 and nan as a result. In this case, it means that there is not enough information to predict the score, and for this lesson, you will leave it as is.

<!-- grade_id: cell-3d0537bb94c333e7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
a =  ~np.isnan(np.expand_dims(ratings, 0)) * np.expand_dims(similarities, 2) 
weights = a / np.expand_dims(np.sum(a, axis=1), 1)
```

Output:
```
/tmp/ipykernel_859/3578830350.py:3: RuntimeWarning: invalid value encountered in divide
  weights = a / np.expand_dims(np.sum(a, axis=1), 1)
```

```python
print(a.shape)
```

Output:
```
(32, 943, 1682)
```

<!-- grade_id: cell-04b8a47e95e8dd09, points: 2, grade: true, solution: false -->
```python
# TEST
assert weights.shape == (32, 943, 1682)
assert np.sum(np.isnan(weights)) == 31119
```

**Step 3**: For each user in the batch, calculate scores (for all the moves) by multiplying the score of that user by the corresponding weight in `weight` array. Then, store the result into a line of the array `filled_ratings`.

<!-- grade_id: cell-e73c0a38ffa925db, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
filled_ratings = np.nansum(weights * ratings, axis=1)
```

<!-- grade_id: cell-9a094032ccb1726a, points: 1.5, grade: true, solution: false -->
```python
# TEST
filled_batch = filled_ratings[start:end]
filled_nanvals = filled_batch[np.isnan(ratings[start:end])]
assert np.array_equal(filled_nanvals[:3].round(1),
                      np.array([3.6, 3.5, 4. ]))
assert np.array_equal(filled_nanvals[-3:].round(1),
                      np.array([2., 3., 3. ]))
```

For now, you can run on all the data (0.5p).

<!-- grade_id: cell-9ee6c3bd8e2dea97, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def predict(start):
    batch = ratings[start:start + batch_size]
    similarities = np.nanmean(np.abs(np.expand_dims(ratings, 0) - np.expand_dims(batch, 1)), axis=2)
    similarities = 1 / (similarities + 0.001)
    similarities[np.isnan(similarities)] = 0.0
    
    weights = ~np.isnan(np.expand_dims(ratings, 0)) * np.expand_dims(similarities, 2)
    weights = weights / np.expand_dims(np.sum(weights, axis=1), 1)
    
    return np.nansum(weights * ratings, axis=1)

filled_ratings = np.empty_like(ratings)
for start in range(0, ratings.shape[0], batch_size):
    filled_ratings[start:start + batch_size,:] = predict(start)
```

Output:
```
/tmp/ipykernel_859/346904947.py:4: RuntimeWarning: Mean of empty slice
  similarities = np.nanmean(np.abs(np.expand_dims(ratings, 0) - np.expand_dims(batch, 1)), axis=2)
/tmp/ipykernel_859/346904947.py:9: RuntimeWarning: invalid value encountered in divide
  weights = weights / np.expand_dims(np.sum(weights, axis=1), 1)
```

<!-- grade_id: cell-5e6aaf23d47e5193, points: 0.5, grade: true, solution: false -->
```python
# TEST
filled_nanvals = filled_ratings[np.isnan(ratings)]
assert np.array_equal(filled_nanvals[:3].round(1),
                      np.array([3.6, 3.5, 4. ]))
assert np.array_equal(filled_nanvals[-3:].round(1),
                      np.array([0., 3., 3. ]))
```